Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE"

---

## 1 
Find the last element of a list. <br>
Eg., l = [1, 2, 3, 4] <br>
last_elemnt(l) -> 4

<!-- grade_id: last_element, points: 0, grade: false, solution: true -->
```python
def last_element(l):
    """
    Returns last element of list l
    """
    # YOUR CODE HERE
    return l[-1]
    
```

Your function should return 4. Check that it does

```python
last_element([1, 2, 3, 4])
```

Output:
```
4
```

<!-- grade_id: correct_last_element, points: 1, grade: true, solution: false -->
```python
"""Testing code for last_element"""
assert last_element([1]) == 1
assert last_element(["Hello", "World"]) == "World"

```

## 2 
Find the number of elements in a list.<br> 
Eg., l = [1, 2, 3, 4] <br>
num_elemnts(l) -> 4

<!-- grade_id: num_elements, points: 0, grade: false, solution: true -->
```python
def num_elements(l):
    """
    Returns number of elements in the list l
    """
    # YOUR CODE HERE
    return len(l)
    
```

<!-- grade_id: correct_num_elements, points: 1, grade: true, solution: false -->
```python
"""Testing code for num_elements"""
assert num_elements([2, 3, 4])==3

```

## 3 
Reverse a list and return it. <br>
Eg., l = [1, 2, 3, 4] <br>
reverse_list(l) -> [4, 3, 2, 1]

<!-- grade_id: reverse_list, points: 0, grade: false, solution: true -->
```python
def reverse_list(l):
    """
    Returns the reverse of a list l
    """
    # YOUR CODE HERE
    return l[::-1]

    
```

<!-- grade_id: correct_reverse_list, points: 0, grade: true, solution: false -->
```python
"""Testing code for reverse_list"""
assert reverse_list([2, 3, 4])== [4, 3, 2]

```

## 4 
Find whether a list a palindrome -> a sequence that reads the same forwards and backwards <br>
is_palindrome([1, 2, 1]) -> True <br>
is_palindrome([1, 2, 3, 2, 1] -> True <br>
is_palindrome([1, 2, 3, 1]) -> False

<!-- grade_id: is_palindrome, points: 0, grade: false, solution: true -->
```python
def is_palindrome(l):
    """
    Returns True if l is a palindrome, False otherwise
    """
    # YOUR CODE HERE
    rev_list=l[::-1]
    return rev_list==l
    

```

<!-- grade_id: correct_is_palindrome, points: 2, grade: true, solution: false -->
```python
"""Testing code for is_palindrome"""

assert(is_palindrome([1, 2, 1])) == True
assert(is_palindrome([1, 2, 3, 2, 1])) == True
assert(is_palindrome([1, 2, 3, 4])) == False

```

### 5
Replace consecutive duplicate elements of list with single element. <br>
Eg., l = [a, a, a, b, b, c, a, a, d, d, d, x, x] <br>
compress(l) -> [a, b, c, a, d, x]

<!-- grade_id: compress, points: 0, grade: false, solution: true -->
```python
def compress(l):
    """
    Returns a list with consecutive duplicate elements replaced by a single element
    """
    # YOUR CODE HERE
    comp=[]
    comp+=[l[0]]
    for i in l:
        if(comp [-1] != i):
            comp+=[i]
    return comp

```

<!-- grade_id: correct_compress, points: 3, grade: true, solution: false -->
```python
"""Testing code for compress"""

assert(compress([1, 2, 2])) == [1, 2]
assert(compress([1, 2, 2, 2, 1, 1, 3, 'x', 'x', 'x'])) == [1, 2, 1, 3, 'x']

```

## 6 
Pack consecutive duplicates of list elements into sublists. <br>
Eg., l = [1, 1, 1, 2, 2, 3, 3, 4]
pack(l) -> [[1,1,1], [2,2], [3,3], [4]]

<!-- grade_id: pack, points: 0, grade: false, solution: true -->
```python
def pack(li):
    """
    Returns a list with consecutive duplicate elements packed into sublists
    """
    # YOUR CODE HERE
    finallist=[]
    templist=[li[0]]
    for i in li[1:]:
        if i == templist[-1]:
            templist+=[i]
        else:
            finallist+=[templist]
            templist=[i]
    if(len(templist)!=0):
        finallist+=[templist]
    print(finallist)
    return finallist
```

<!-- grade_id: correct_pack, points: 0, grade: true, solution: false -->
```python
"""Testing code for pack"""
assert(pack([1, 1, 1, 2]) == [[1, 1, 1], [2]])
assert(pack([1, 1, 1, 2, 1, 1, 3, 3, 3])) == [[1, 1, 1], [2], [1, 1], [3, 3, 3]]
```

Output:
```
[[1, 1, 1], [2]]
[[1, 1, 1], [2], [1, 1], [3, 3, 3]]
```

## 7
Given two indices, i and k, the slice is the list containing the elements between the i'th and k'th element of the original list (left limit included). Start counting the elements with 0.

Eg., l = [1, 3, 9, 8, 7];
slice(l, 1, 3) --> [3, 9]


<!-- grade_id: slice, points: 0, grade: false, solution: true -->
```python
def slice(l, i, k):
    """
    Returns a list containing the elements between i'th and k'th elements of original list l.
    """  
    # YOUR CODE HERE
    li=[]
    if(k<=len(l)and i>=0):
      for i in range(i,k):
        li+=[l[i]]
    elif( i>=0 and k>len(l)):
      for i in range(i,len(l)):
        li+=[l[i]]
        
    # note you may use return l[i:k] but remember to perform necessary check for edge cases
    return li
     
```

<!-- grade_id: correct_slice, points: 1, grade: true, solution: false -->
```python
"""Testing code for slice"""
assert(slice([1, 3, 8, 9, 7], 1, 3)) == [3,8]
assert(slice([1, 4, 6, 'x', 9, 0], 2, 10)) == [6, 'x', 9, 0]

```

## 8 
Given a list l, index i and element elem, return a new list with elem at index i <br>
Eg., l = [1, 3, 9, 8, 7]; <br>
insert_element(l, 1, 3) --> [1, 3, 9, 1, 8, 7]

<!-- grade_id: insert_element, points: 0, grade: false, solution: true -->
```python
def insert_element(l, i, elem):
    """
    Returns a new list containing elem at index i. If i > len (l), insert element at the end of the list
    """      
    # YOUR CODE HERE
    li=[]
    for element in l:
      if(i<0):
        li=l
        break
      if(i>=len(l)):
        li=l+[elem]
        break
      if(i<len(l)):
        li[i:]=[elem]+l[i:]
    return li
        
      
    
```

<!-- grade_id: correct_insert_element, points: 1, grade: true, solution: false -->
```python
"""Testing code for insert_element"""
assert(insert_element([1, 2, 3, 4,], 2, 5)[2]) == 5
assert(insert_element([1, 5, ], 3, 5)) == [1, 5, 5]

```

```python

```